---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## 컨벤션에 대해서

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

(각 2점)

식별자 이름의 형태를 통해 직관적으로 이것이 함수인지, 변수인지 혹은 클래스명을 바로 알 수 있습니다.

저희가 아직 실질적으로 함수, 클래스를 만드는 방법을 배우지는 않았지만, 미리 컨벤션 규칙을 익히려 합니다.

`@instructions`
아래 단어를 활용하여 변수명, 함수명, 클래스명을 컨벤션 규칙에 따라 작성해주세요.

**_customer list_**

`@hint`
띄워쓰기 불가 / 숫자로 시작 불가 / 한글 불가 / 이미 선언된 함수명은 피하세요

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
# customer list

# (1) 아래에 NAME 를 지우고 위 단어로 컨벤션에 맞게 함수명을 작성해주세요.

def NAME():
  pass

# (2) 아래에 %%%% 를 지우고 위 단어로 컨벤션에 맞게 변수명을 입력해주세요.

NAME = 900

# (3) 아래에 %%%% 를 지우고 위 단어로 컨벤션에 맞게 클래스명을 입력해주세요.
# class 아래에 def, self 등 코드는 지금은 신경쓰지 않으셔도 됩니다.

class NAME(object):
  def __init__(self):
    self.name = "Hello"
```

`@solution`
```{python}
# (1) 아래에 NAME 를 지우고 위 단어로 컨벤션에 맞게 함수명을 작성해주세요.

def customer_list():
  pass

# (2) 아래에 NAME 를 지우고 위 단어로 컨벤션에 맞게 변수명을 입력해주세요.

customer_list = 900

# (3) 아래에 NAME 를 지우고 위 단어로 컨벤션에 맞게 클래스명을 입력해주세요.
# class 아래에 def, self 등 코드는 지금은 신경쓰지 않으셔도 됩니다.

class customerList(object):
  def __init__(self):
    self.name = "Hello"
```

`@sct`
```{python}
#SUBMISSION CORRECTNESS TESTS
```

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 993edf8798
xp: 50
```

만약 그런거라면 어떻게 해야할까요?

`@hint`
gogo

`@possible_answers`
haha

`@feedback`
ggg

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 31e17b2a6f
xp: 50
```

`@projector_key`
c562b68e7ed7acf0428e0122a43bf7a5

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 52e62b140b
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
