---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

이 것은 문제입니다. A를 통해서 B를 구하세요. 화이팅입니다.

`@instructions`
이 것은 가이드입니다. 이렇게 저렇게 풀어보세요. 생각보다 어렵습니다.

`@hint`
이 것은 힌트입니다.

`@pre_exercise_code`
```{python}
print("helloworld")
```

`@sample_code`
```{python}

```

`@solution`
```{python}
print("worldhello")
```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 993edf8798
xp: 50
```

만약 그런거라면 어떻게 해야할까요?

`@hint`
gogo

`@possible_answers`
haha

`@feedback`
ggg

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 31e17b2a6f
xp: 50
```

`@projector_key`
c562b68e7ed7acf0428e0122a43bf7a5

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 52e62b140b
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
